---
title: "R4ScHARR"
tutorial:
  id: "R4ScHARR.day.1"
  version: 0.1.4
output: 
  learnr::tutorial:
    progressive: false
    allow_skip: true
    theme: default
    css: "css/mystyle.css"
runtime: shiny_prerendered

---

<!--ex-setup-r.Rmd-->

```{r setup, include=FALSE}

# devtools::install_github("rstudio/learnr") # use dev. version?
library(learnr)
# devtools::install_github("rstudio-education/gradethis")
library(gradethis)

tutorial_options(exercise.timelimit = 5, exercise.checker = gradethis::grade_learnr)
knitr::opts_chunk$set(echo = FALSE)
source("./learnr_bugfix.R")

# create some global objects for the exercises
.hidden_dat <- read.csv("https://raw.githubusercontent.com/ScHARR-PHEDS/R4ScHARR/master/data/df_framingham.csv")
.hidden_cleaned_dat <- .hidden_dat[!is.na(.hidden_dat$sysBP),]  
.hidden_dat_female <- .hidden_cleaned_dat[.hidden_cleaned_dat$sex == "female",]
.hidden_dat_male <- .hidden_cleaned_dat[.hidden_cleaned_dat$sex == "male",]

```

<!-- 
<img src="www/banner.png" width="100%"  align="left"> 
-->
<img src="www/banner_small.png" width="100%"  align="left"> 

<br><br>

## Welcome

<font size="+1"> **Welcome to the R-4-ScHARR introductory R course site**</font>

<br>

**This site hosts:**

  * Course schedule
  * Instructions on installing R and R Studio 
  * Exercises for sessions 1-3
  
  <br>
  
**Links to other course material:**

* <a href="https://scharr-pheds.github.io/R4ScHARR/" target="_blank">Course book and code examples</a>
* <a href="https://drive.google.com/file/d/1CYWJPdE0wp-WjtAmcL2EmskhjBlhah-x/view?usp=sharing" target="_blank">Presentation slides</a>
* <a href="https://github.com/ScHARR-PHEDS/R4ScHARR" target="_blank">ScHARR Github Repository</a>

We hope that you find the content useful.

*Rob & Paul*

<br>



****

## About

<br>

This course is an introduction to programming in R. It is designed to provide delegates with a basic working knowledge of R. We move through the material very slowly, focusing on the 'base R' methods in the knowledge that many of the delegates will go on to learn alternative methods (e.g. pipes). If you have previous programming experience this course is likely to seem slow. 

If you have specific questions relating to the use of R for a project, Rob & Paul may be able to help with: 
  
  * Epidemiology/medical statistics
  * Micro-simulation models (Health Economics)
  * Markov models (Health Economics)
  * Cost Effectiveness Modelling in R (Health Economics)
  * Capacity modeling (Health Economics)
  * Web-based applications (R-Shiny)
  * Geo-spacial analysis
  * Data scraping & APIs
  * Creating beautiful plots & visualizations.    




<br><br>

### Who we are

```{r echo=FALSE}
# Rob profile
fluidRow(
         div(style="display: inline-block;vertical-align:top;",
             img(src="www/rob.png",width="200px", align="left")
             # HTML('<img src="www/rob.png" width=200px />')
             ),
         div(style="display: inline-block;vertical-align:top; width: 45%;",HTML('<a href="https://www.linkedin.com/in/robert-smith-53b28438/" target="_blank">Robert Smith</a> joined ScHARR in 2016. His research focuses on the methods used to estimate the costs and benefits of public health interventions, with a specific interest in microsimulation modelling in R. He has become increasingly interested in the use of web based user interfaces (in R-Shiny) to make research more transparent and to aid decision makers. While doing his PhD, Robert has been involved in multiple projects with the World Health Organization and parkrunUK.')),
         
         # space
         br(),br(),br(),
          
         
# Paul profile
         div(style="display: inline-block;vertical-align:top; width: 45%; padding-left:20px;",HTML('<a href="https://www.sheffield.ac.uk/scharr/staff-pgrs/studentprofiles/paulschneider" target="_blank">Paul Schneider</a> joined ScHARR in 2018. He is working on conceptual and methodological problems in valuing health outcomes in economic evaluations. A medical doctor and epidemiologist by training, he has used R in various research projects, ranging from the modeling costs of breast cancer, and value of information analyses, to the monitoring of influenza in real-time using online data. He is a keen advocate of open science practices.')),
         div(style="display: inline-block;vertical-align:top;",
             img(src="www/paul.jpg",width="200px", align="right")
             )
)
             
         




```


## Course schedule


> Before the start of the course, please make sure that you have **R** and **R Studio** on your computer. The tutorial on the next page shows you how to install the software, if you haven't done this already. <br>
Go to: [How to install R](#link1)


<iframe src="https://docs.google.com/presentation/d/e/2PACX-1vTa0OgQxADt6OKoO2svUcoCt062PlxKy_zW7gvYQ9IXr9M1ogf--4E21sXO5fctI7w9jkuuQyqDaWGJ/embed?start=false&loop=false&delayms=60000&rm=minimal" frameborder="1"  allowfullscreen="false" mozallowfullscreen="false" webkitallowfullscreen="false" height="480" width="628" id="iframe1"></iframe>


<br>

****


<!-- create button to link to next page 
<a href="exercises.Rmd#section-link1" class="btn btn-default skip">Go to install R guide</a>
-->

<!---------- How to install R ------------->


## How to install R{#link1}

R is a free software environment for statistical analysis and data science. 

There are two pieces of software you should install to work with R: first, base R itself, and then R Studio, which is is an almost universally used user interface - a so called integrated development environment (IDE) - that makes working with R much easier.

R is made freely available by the Comprehensive R archive Network (CRAN), in the UK it can be downloaded through the University of Bristol *[here](https://www.stats.bris.ac.uk/R/)*. There are options of downloading R for Linux, Mac and Windows.

You should download the free 'RStudio Desktop Open Source Licence' version for the laptop you will be attending the course with from *[RStudio.com](https://www.rstudio.com/products/rstudio/#Desktop)*. If you have time before the course, it would be hugely beneficial to get familiar with RStudio.

The videos below provide a step by step guide to install R and R Studio on your computer.

<br><br>

### How to install base R

![](https://vimeo.com/203516510)

<br><br>

### Install RStudio

![](https://vimeo.com/203516968)

<br>

<!-- create button to link to next page 
<a href="exercises.Rmd#section-instructions" class="btn btn-default skip">Continue</a>
-->

If you are still struggling try this link:
[Installing-R](https://rstudio-education.github.io/hopr/starting.html)


<!---------- INSTRUCTIONS ------------->




## Instructions{#instructions}

When learning a programming language, it is important to *use it* and get a feel for how things work. In a in-person meeting, we would have multiple practical exercises that involve typing-in code and solving some problems with R/R Studio. We would walk around the room and answer any questions, in case you have technical problems, encounter cryptic error messages, or get stuck otherwise - this is part of the normal process of learning R and the initial confusion will soon dissolve.

However, in a virtual meeting, practical exercises are a bit of a challenge: it is more difficult to provide 1:1 support and following the course and using R at the same time might not be possible, if both application windows have to share the same screen. Therefore, practical examples are not conducted in R itself, but on the next pages of this website. You will there find code blocks, that allow you to enter and run your own R code.

This allows you to write R code in a safe environment, with automatic feedback and support. However, we are still available to answer any questions you may have as you move through the material.

<br><br>

<div class="box">

#### First taste

****

> Now, give it a try: replace `"YOUR_NAME"` in the code block below with your own name and then click <span style="display: inline;">
  <a class="btn btn-success btn-xs btn-tutorial-run" type="button" title="" data-icon="fa-play"><i class="fa fa-play"></i> Run Code</a> </span>.


<div class="alert alert-info">
**Info**:  The code is evaluated (run) and then prints the resulting output as it would if you ran this code within R. If you get stuck, click <a class="btn btn-light btn-xs btn-tutorial-solution" role="button" title="Hint"><i class="fa fa-lightbulb-o"></i> Hint</a> to get a hint, or <span style="display: inline;">
  <a class="btn btn-light btn-xs btn-tutorial-solution" role="button" title="Start Over"><i class="fa fa-refresh"></i> Start Over</a> to revert and restore the original code.
</div>
</span> 
<div class="alert alert-warning">
**Important!** Do not click run or submit the code if it is completely empty. This causes an error and disables the block of code.
</div>



<br>

```{r my_first_lines_of_R, exercise=TRUE,exercise.cap="My first R code"}
id = "YOUR_NAME"
cat("Hello",id,"- welcome to R4ScHARR!")
```


```{r my_first_lines_of_R-hint}
# try:
id = "Alice"
cat("Hello",id,"- welcome to R4ScHARR!")
```

</div>

<br>

<!-- create button to link to next page -->
<a href="exercises.Rmd#section-session1" class="btn btn-default skip">Start Session 1 Exercise</a>






<!---------- SESSION 1 ------------->

## Session 1 exercises{#session1}

<br>

### R as a fancy calculator

To get started, let's use R to solve some simple computations. We give you the code for the first exercise, and you have to find the solution for the others.


<br><br>

<div class="box">

#### Exercise 1.1 

****

> What is 1 plus 1?


<div class="alert alert-info">
**Info**: Just click on the <span style="display: inline;"> <a class="btn btn-success btn-xs btn-tutorial-run" type="button" title="" data-icon="fa-play"><i class="fa fa-play"></i> Run Code</a> </span> button on the right to get the result.
</div>

```{r code1_1, exercise=TRUE, exercise.eval=FALSE}
1+1
```

```{r code1_1-check}
grade_result(
  pass_if(~ identical(.result, 2), "That's correct!")
)
```

</div>

<br><br>


<div class="box">
#### Exercise 1.2 - Now, it's your turn!

****

> What is 3 times 17 ?

<div class="alert alert-info">
**Info**:

  1. Type in the code to compute the results below
  2. Click <span style="display: inline;"> <a class="btn btn-success btn-xs btn-tutorial-run" type="button" title="" data-icon="fa-play"><i class="fa fa-play"></i> Run Code</a> </span>
  3. If you are unsure, click on <a class="btn btn-light btn-xs btn-tutorial-solution" role="button" title="Hint"><i class="fa fa-lightbulb-o"></i> Hint</a>
  4. To check if your answer is correct, click on <a class="btn btn-primary btn-xs btn-tutorial-run" type="button" title="Submit Answer" data-check="1" data-icon="fa-check-square-o"><i class="fa fa-check-square-o"></i> Submit Answer</a>
</div>
<br>

```{r code1_2, exercise=TRUE, exercise.eval=FALSE}

```

```{r code1_2-hint}
# Try:
  3 * 17
```

```{r code1_2-check}
grade_result(
  pass_if(~ identical(.result, 51), "Good job!"),
  incorrect = "That's not quite right. Try again or click on 'Hint'.",
  glue_incorrect = "That's not quite right. Try again or click on 'Hint'.",
  glue_correct = "Good job!"
)
```

</div>

<br>

<div class="box">
#### Exercise 1.3

****

> Create an object `d` equal to 10.

<div class="alert alert-info">
**Info**: this does not give you any output.
</div>

```{r code1_3a, exercise=TRUE, exercise.eval=FALSE}

```

```{r code1_3a-solution}
d = 10
```

```{r code1_3a-check}

grade_code( 
  glue_incorrect = "That's not quite right. Try again or click on 'Solution'."
  )


#   grade_result(
#   pass_if(~ identical(.result, 51), "Good job!"),
#   incorrect = "That's not quite right. Try again or click on 'Hint'.",
#   glue_incorrect = "That's not quite right. Try again or click on 'Hint'.",
#   glue_correct = "Good job!"
# )
```
</div>
<br>

<div class="box">
#### Exercise 1.4

****

> Divide `d` by 5. 

```{r code1_4-setup}
d = 10
```

```{r code1_4, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code1_4-setup" }

```

```{r code1_4-solution}
d/5
```

```{r code1_4-check}
grade_result(
  pass_if(~ identical(.result, 2), "Well done!")
)
```

</div>
<br>

<div class="box">
#### Exercise 1.5

****

> Overwrite `d`, let it now be equal to 20, and 'print' `d`, to see the results


```{r code1_5, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code1_4-setup" }

```

```{r code1_5-solution}
d = 20
d
```

```{r code1_5-check}
grade_result( 
  pass_if(~ identical(.result, 20), "Great!"),
  glue_incorrect = "That's not quite right. Try again or click on 'Solution'."
  )
```

</div>
<br>

<div class="box">
#### Exercise 1.6

****

> Overwrite `d` (which is now 20) with 4 times `d`, and print `d` top see the result

```{r code1_6-setup}
d = 20
```

```{r code1_6, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code1_6-setup"}

```

```{r code1_6-solution}
d = d * 4
d
```

```{r code1_6-check}
grade_result( 
  pass_if(~ identical(.result, 80), ""),
  glue_incorrect = "That's not quite right. Try again or click on 'Solution'."
  )
```


</div>
<br>

<div class="box">
#### Exercise 1.7

****

> Overwrite `d` (which is now 80), with `d` times `d`, and print the result

```{r code1_7-setup}
d = 80
```

```{r code1_7, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code1_7-setup" }

```

```{r code1_7-solution}
d = d * d
d
```

```{r code1_7-check}
grade_result( 
  pass_if(~ identical(.result, 6400), ""),
  glue_incorrect = "That's not quite right. Try again or click on 'Solution'."
  )
```

</div>
<br>

<div class="box">
#### Exercise 1.8

****

> Create a new object `A` and set it equal to 3

```{r code1_8-setup}
d = 6400
```

```{r code1_8, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code1_8-setup" }

```

```{r code1_8-solution}
A = 3
```

```{r code1_8-check}
grade_result( 
  pass_if(~ identical(.result, 3), ""),
  glue_incorrect = "That's not quite right. Try again or click on 'Solution'."
  )
```


</div>
<br>

<div class="box">
#### Exercise 1.9

****

> Create another new object `res`, let it be the product of `d` and `A`, and output the result

<div class="alert alert-info">
**Info**: R is case-sensitive. If you use `a` instead if `A`, it will give you an error message - try it for yourself."

</div>
<br>

```{r code1_9-setup}
d = 6400
A = 3
```

```{r code1_9, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code1_9-setup" }

```

```{r code1_9-solution}
res = d*A
```

```{r code1_9-check}
grade_result( 
  pass_if(~ identical(.result, 19200), ""),
  glue_incorrect = "That's not quite right. Try again or click on 'Solution'."
  )
```

</div>
<br>

<div class="box">
#### Exercise 1.10

****

> Finally, take the average of `d`, `A` and `res`

```{r code1_10-setup}
d = 6400
A = 3
res = d * A
```

```{r code1_10, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code1_10-setup" }

```

```{r code1_10-solution}
(d + A + res) / 3
# you could also use
mean(c(d, A, res))
```

```{r code1_10-check}
grade_result(
  pass_if(~ identical(.result, (6400 + 3 + 6400*3) / 3), "")
)
```

</div>

<br>

****


### Evaluations

Now, use R to evaluate the following equations for you:

<br>

<div class="box">
#### Exercise 1.11

****

> Is 4 greater than 2?  


```{r code1_11, exercise=TRUE, exercise.eval=FALSE}

```

```{r code1_11-solution}
4 > 2
```

```{r code1_11-check}
grade_result(
  pass_if(~ identical(.result, T), "")
)
```

</div>
<br>


<div class="box">
#### Exercise 1.12

****

> Is 5 less than 3?  


```{r code1_12, exercise=TRUE, exercise.eval=FALSE}

```

```{r code1_12-solution}
5<3
```

```{r code1_12-check}
grade_result(
  pass_if(~ identical(.result, F), "")
)
```

</div>
<br>

<div class="box">
#### Exercise 1.13

****

> Is 6.2 equal to 12.4/2?   


<div class="alert alert-info">
**Hint**: use `==` not `=`   
</div>

```{r code1_13, exercise=TRUE, exercise.eval=FALSE}

```

```{r code1_13-solution}
6.2 == 12.4/2
```

```{r code1_13-check}
grade_result(
  pass_if(~ identical(.result, T), "")
)
```

</div>
<br>


<div class="box">
#### Exercise 1.14

****

>Is 5 equal to or greater than 4? 

<div class="alert alert-info">
**Hint**: use `>=`   
</div>

```{r code1_14, exercise=TRUE, exercise.eval=FALSE}

```


```{r code1_14-solution}
5 >= 4
```

```{r code1_14-check}
grade_result(
  pass_if(~ identical(.result, T), "")
)
```

</div>
<br>



<div class="box">
#### Exercise 1.15* - last one for this session

****

> I   Create an object `foo` and set it to be 8 <br>
> II  Create another object `bar`, and set it to be 7 <br>
> III Is `foo` divided by `bar`  greater than 1.14 ?

<div class="alert alert-info">
**Info**: Put the code for subtask I in line 1, the code for subtask II in line 2,the code for subtask III in line 3, and hit <span style="display: inline;">
  <a class="btn btn-success btn-xs btn-tutorial-run" type="button" title="" data-icon="fa-play"><i class="fa fa-play"></i> Run Code</a> </span>.
</div>

```{r code1_15-setup}
foo<-8
bar<-7
```

```{r code1_15, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code1_15-setup" }

```


```{r code1_15-solution}
8/7 > 1.14
```

```{r code1_15-check}
grade_result(
  pass_if(~ identical(.result, F), "")
)
```


****

</div>
<br>


<!-- 

  ADDITIONAL EXERCISES 
  
  2a) Create an object m equal to 7. 
  
  b) Overwrite m with m = m times 10. 
  
  c) Create an object p equal to 2. 
  
  d)Overwrite p with p = p times 12. 
  
  e) Create an object w equal to m divided by p. 
  
  f) What values do m, p and w take? 
  ```{r,echo=FALSE, eval=FALSE}
m <- 7 ; m <- m*10 ; p <- 2 ; p <- p*12 ; w <- m/p
```
-->


****

<!-- create button to link to next page -->
<a href="exercises.Rmd#section-session2" class="btn btn-default skip">Start Session 2 Exercise</a>



## Session 2 exercises{#session2}

<br>


```{r include=FALSE}
dat = read.table(text=
"height weight first_name sex      bmi
1   1.38     31      Alice   F 16.27809
2   1.45     35        Bob   M 16.64685
3   1.21     28      Harry   M 19.12438
4   1.56     40       Jane   F 16.43655")
primes <- c(2,3,5,7,11) 
```


### Subsetting

Suppose we want to subset this data frame, called `dat`:

```
> dat
 height weight first_name sex      bmi
1   1.38     31      Alice   F 16.27809
2   1.45     35        Bob   M 16.64685
3   1.21     28      Harry   M 19.12438
4   1.56     40       Jane   F 16.43655
```



<div class="box">
#### Exercise 2.1
    
****
    
> Select the **3rd row** from the data frame `dat`
    

    
```{r code2_1-setup}
dat = read.table(text=
"height weight first_name sex      bmi
1   1.38     31      Alice   F 16.27809
2   1.45     35        Bob   M 16.64685
3   1.21     28      Harry   M 19.12438
4   1.56     40       Jane   F 16.43655")
primes <- c(2,3,5,7,11) 
```
  
```{r code2_1, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code2_1-setup" }
  
```
  
```{r code2_1-solution}
  dat[3,]
```
  
```{r code2_1-check}
grade_result(
  pass_if(~ identical(.result, dat[3,]), "")
  )
```

</div>
<br>

<div class="box">
#### Exercise 2.2
    
****
    
> Select the `weight` variable from the data frame using your preferred method.  

<div class="alert alert-info">
**Info**: Remember, there are multiple ways to do this: you can use the `$` to index the right column, or use the `[ , ]`  with a number or a  "variable_name".
</div>
<br>    

    
```{r code2_2-setup}
dat <- dat
```
  
```{r code2_2, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code2_2-setup" }
  
```
  
```{r code2_2-solution}
  dat$weight 
# or dat[,2] 
# or dat[,"weight"]
```
  
```{r code2_2-check}
grade_result(
  pass_if(~ identical(.result, dat$weight ), "")
  )
```
  
</div>
<br>
  
  
  
<div class="box">
#### Exercise 2.3
    
****
    
> Select `Alice`'s data from the data frame.   

    
<div class="alert alert-info">
**Info**: It might be helpful to type in `dat` and press <span style="display: inline;">
  <a class="btn btn-success btn-xs btn-tutorial-run" type="button" title="" data-icon="fa-play"><i class="fa fa-play"></i> Run Code</a> </span>. to check which row Alice's data is in.
</div>
<br>
    
```{r code2_3-setup}

```
  
```{r code2_3, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code2_3-setup" }
  
```
  
```{r code2_3-solution}
dat[dat$first_name == "Alice",]
# or simply:
# dat[1,]
```
  
```{r code2_3-check}
grade_result(
  pass_if(~ identical(.result, dat[1,]), "")
  )
```
  
</div>
<br>
  


<div class="box">
#### Exercise 2.4
    
****
    
> Print `dat` without it's first row
    
<div class="alert alert-info">
**Info**: You can select rows 2,3, and 4, but you can also show everything except row 1 - try the latter approach.
</div>
<br>
    
```{r code2_4-setup}
dat <- dat
```
  
```{r code2_4, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code2_4-setup" }
  
```
  
```{r code2_4-solution}
dat[-1,]
# but you could also use:
# dat[2:4,]
# or
# dat[ c(F, T, T, T), ]
```
  
```{r code2_4-check}
grade_result(
  pass_if(~ identical(.result, dat[-1,]), "")
  )
```
  
</div>
<br>



<div class="box">
#### Exercise 2.5
    
****
    
> Subset the data frame to show just the data for the females
    
<div class="alert alert-info">
**Info**: Remember, `F` (=FALSE) and `"F"` (= F as a character) have different meanings in R.
</div>
<br>
    
```{r code2_5-setup}

```
  
```{r code2_5, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code2_5-setup" }
  
```
  
```{r code2_5-solution}
  dat[dat$sex == "F",]
```
  
```{r code2_5-check}
grade_result(
  pass_if(~ identical(.result, dat[dat$sex == "F",]), "")
  )
```
  
</div>
<br>  



<div class="box">
#### Exercise 2.6
    
****
    
> Create a vector `primes` with elements `2`,`3`,`5`,`7`, and `11`.
    
<div class="alert alert-info">
**Info**: You can `c`ombine elements into a vector with the `c(element_1, element_2,...)` command.
</div>
<br>
    
```{r code2_6-setup}

```
  
```{r code2_6, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code2_6-setup" }
  
```
  
```{r code2_6-solution}
primes <- c(2,3,5,7,11)  
```
  
```{r code2_6-check}
grade_code()
```
  
</div>
<br>



<div class="box">
#### Exercise 2.7
    
****
    
> Print all `primes` that are smaller than or equal to `5` using subsetting.
    
    
```{r code2_7-setup}
primes <- c(2,3,5,7,11)  
```
  
```{r code2_7, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code2_7-setup" }
  
```
  
```{r code2_7-solution}
primes[primes <= 5]
```
  
```{r code2_7-check}
grade_result(
  pass_if(~ identical(.result, primes[primes <= 5]), "")
  )
```
  
</div>
<br>



<div class="box">
#### Exercise 2.8* - last one for this session
    
****
    
> Create an 'animal top speed data frame': `speed_dat`. It should have three columns, named `animal`, `speed`, `flying`, and should contain the following information:
<br> <br> The `"lion"` can run `80`km/h, it's not flying. <br> 
The `"marlin"` can swim `129`km/h, it's also not flying. <br>
Finally, the `"eagle"` can do `240`km/h, and of course it's flying.

<div class="alert alert-info">
**Info 1**: The vector `flying` should be set to `TRUE` or `FALSE`. 
</div>
<br>
    
<div class="alert alert-info">
**Info 2**: Use `data.frame(vector_1, vector_2,...)` to combine multiple vectors into a data frame.
</div>

    
```{r code2_8-setup}

```
  
```{r code2_8, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code2_8-setup" }
  
```
  
```{r code2_8-solution}
animal <- c("lion", "marlin", "eagle")  
speed <- c(80, 129, 240)
flying <- c(F, F, T)
speed_dat <- data.frame(animal, speed, flying)
speed_dat
```
  
```
This task is a bit to complicated to have it checked automatically - please compare your result with the solution.
```
  
</div>
<br>


****


<!-- create button to link to next page -->
<a href="exercises.Rmd#section-session3" class="btn btn-default skip">Start Session 3 Exercise</a>









<!------------     SESSION 3 --------------------->










## Session 3 exercises{#session3}

<br>

### Doing data analysis in R

This exercise takes you through some of the main steps for analysing a data set in R.
  
  
<br>
  
<div class="box">
#### Exercise 3.1
    
****
    
> Use `ls()` to check your global enviroment.
    
<div class="alert alert-info">
**Info**: When starting a new project, you should check your 'global environment' (= the objects existing in your R session) and make sure you don't have old, useless objects hanging around.
</div>
<br>
    
```{r code3_1-setup}
some_old_forgotten_thing <- c("I am a useless object", "Someone has forgotten to remove me from the enviroment","Delete me!")
```
  
```{r code3_1, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_1-setup" }
  
```
  
```{r code3_1-solution}
ls()
```
  
```{r code3_1-check}
grade_code()
```
  
</div>
<br>



  
<div class="box">
#### Exercise 3.2
    
****
    
> Clear the enviroment using the `rm()` function and check if it has worked.
    
<div class="alert alert-info">
**Info**:You can use `rm` either with varibales' names (`c("varname_1", "varname_2",...)`), or, to remove all objects, you can use `ls()` itself as an argument for the function.
</div>
<br>
    
```{r code3_2-setup}
some_old_forgotten_thing <- c("I am a useless object", "Someone has forgotten to remove me from the enviroment","Delete me!")
```
  
```{r code3_2, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_2-setup" }
  
```
  
```{r code3_2-solution}
rm(list=ls())
# or rm("some_old_forgotten_thing")
ls()
```
  
```{r code3_2-check}
grade_result(
  pass_if(~ identical(.result, character()), "")
  )
```
  
</div>
<br>



  
<div class="box">
#### Exercise 3.3
    
****
    
> Load the raw data set into your enviroment and get an overview using the `str()` function.

<div class="alert alert-info">
**Info**: We give you the first line of code, because this is an online exercise and so we have to use a url to load the data. Usually, you would specify a local (relative) path on your computer (e.g. './data/framingham_sample.csv') or select the file manually with `file.choose()`.
</div>
<br>
    
    
```{r code3_3-setup}


```
  
```{r code3_3, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_3-setup" }
dat <- read.csv("https://raw.githubusercontent.com/ScHARR-PHEDS/R4ScHARR/master/data/df_framingham.csv")
```
  
```{r code3_3-solution}
dat <- read.csv("https://raw.githubusercontent.com/ScHARR-PHEDS/R4ScHARR/master/data/df_framingham.csv")
str(dat)
```
  
```{r code3_3-check}
dat <- .hidden_dat
grade_result(
  pass_if(~ identical(.result, str(dat)), "")
  )
```
  
</div>
<br>



  
<div class="box">
#### Exercise 3.4
    
****
    
> Inspect the first few rows of the `dat` data frame 
    
<div class="alert alert-info">
**Info**:`tail(dat)` shows you the last few rows of the data frame.
</div>
<br>
    
```{r code3_4-setup}
dat <- .hidden_dat
```
  
```{r code3_4, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_4-setup" }
  
```
  
```{r code3_4-solution}
head(dat)  
```
  
```{r code3_4-check}
dat <- .hidden_dat
grade_result(
  pass_if(~ identical(.result, head(dat)  ), "")
  )
```
  
</div>
<br>


  
<div class="box">
#### Exercise 3.5
    
****
    
> How many `NA` does `dat$sysBP` have?
    
<div class="alert alert-info">
**Info**: `is.na(vector)` returns `T`/1 if vector is `NA`, and `F`/0 otherwise. 
</div>
<br>
    
```{r code3_5-setup}
dat <- .hidden_dat
```
  
```{r code3_5, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_5-setup" }
  
```
  
```{r code3_5-solution}
  sum(is.na(dat$sysBP) )
```
  
```{r code3_5-check}
dat <- .hidden_dat
grade_result(
  pass_if(~ identical(.result, sum(is.na(dat$sysBP) )), "")
  )
```
  
</div>
<br>


  
<div class="box">
#### Exercise 3.6
    
****
    
> Remove all rows with `NA` from `dat`
    
    
```{r code3_6-setup}
dat <- .hidden_dat
```
  
```{r code3_6, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_6-setup" }
  
```
  
```{r code3_6-solution}
dat <- dat[!is.na(dat$sysBP),]  
```
  
```{r code3_6-check}
dat <- .hidden_dat
dat <- dat[!is.na(dat$sysBP),]
grade_result(pass_if(~ identical(.result, dat), ""))
```
  
</div>
<br>



  
<div class="box">
#### Exercise 3.7
    
****
    
> Create a subset of `dat`, either containing only cases females, in that case name your new dataframe `dat_female`,  or only males, then call your new dataframe `dat_male`.
    
<div class="alert alert-info">
**Info**: Only use one of the two options.
</div>
<br>
    
```{r code3_7-setup}
dat <- .hidden_cleaned_dat
```
  
```{r code3_7, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_7-setup" }

```
  
```{r code3_7-solution}
# Either
dat_female <- dat[dat$sex == "female",]
# or
dat_male <- dat[dat$sex == "male",]
```
  
```{r code3_7-check}
dat <- .hidden_cleaned_dat
dat_female <- dat[dat$sex == "female",]
dat_male <- dat[dat$sex == "male",]

grade_result(
  pass_if(~ identical(.result, dat_female), ""),
  pass_if(~ identical(.result, dat_male), "")
  )
```
  
</div>
<br>




  
<div class="box">
#### Exercise 3.8
    
****
    
> What is the median bloodpressure in your subsetted dataframe?
    
<div class="alert alert-info">
**Info**: Analogous to `mean()`, there is a function in R called `median()`!
</div>
<br>
    
```{r code3_8-setup}
dat_female <- .hidden_dat_female
dat_male <- .hidden_dat_male
```
  
```{r code3_8, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_8-setup" }
  
```
  
```{r code3_8-solution}
# Either 
median(dat_male$sysBP)
# Or
median(dat_female$sysBP)
```
  
```{r code3_8-check}
dat_female <- .hidden_dat_female
dat_male <- .hidden_dat_male

grade_result(
  pass_if(~ identical(.result, median(dat_female$sysBP)), ""),
  pass_if(~ identical(.result, median(dat_male$sysBP)), "")
  )
```
  
</div>
<br>



  
<div class="box">
#### Exercise 3.9
    
****
    
> What is the range (difference between highest and the lowest) of blood pressure values in your sample?
    
<div class="alert alert-info">
**Info**: You may want to use the functions `max()` and `min()` for this.
</div>
<br>
    
```{r code3_9-setup}
dat_female <- .hidden_dat_female
dat_male <- .hidden_dat_male
```
  
```{r code3_9, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_9-setup" }
  
```
  
```{r code3_9-solution}
# Either
max(dat_male$sysBP) - min(dat_male$sysBP) 
# or
max(dat_female$sysBP) - min(dat_female$sysBP) 
```
  
```{r code3_9-check}
dat_female <- .hidden_dat_female
dat_male <- .hidden_dat_male

grade_result(
  pass_if(~ identical(.result, max(dat_male$sysBP) - min(dat_male$sysBP) ), ""),
  pass_if(~ identical(.result, max(dat_female$sysBP) - min(dat_female$sysBP) ), "")
  )
```
  
</div>
<br>



  
<div class="box">
#### Exercise 3.10
    
****
    
> Create a histogram to view the disctribution of `sysBP` in your subsample.
    
<div class="alert alert-info">
**Info**: If you want, you can try to use the `hist(...)` function with an additional style argument: `xlab = "Systolic blood pressure"`.
</div>
<br>
    
```{r code3_10-setup}
dat_female <- .hidden_dat_female
dat_male <- .hidden_dat_male
```
  
```{r code3_10, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_10-setup" }
  
```
  
```{r code3_10-solution}
# Either
hist(dat_female$sysBP, xlab = "Systolic blood pressure")
# or 
hist(dat_male$sysBP, xlab = "Systolic blood pressure")
```
  
```{r code3_10-check}
dat_female <- .hidden_dat_female
dat_male <- .hidden_dat_male

grade_result(
  pass_if(~ identical(.result, hist(dat_male$sysBP)), ""),
  pass_if(~ identical(.result, hist(dat_male$sysBP, xlab = "Systolic blood pressure")), ""),
  pass_if(~ identical(.result, hist(dat_female$sysBP)), ""),
  pass_if(~ identical(.result, hist(dat_female$sysBP, xlab = "Systolic blood pressure")), "")
  )
```
  
</div>
<br>

  
<div class="box">
#### Exercise 3.11
    
****
    
> Fit a linear regression to model the association between `sysBP` and `age`. Call the model either fit_m (for males) or fit_f (for females). Then use `summary()` to inspect the details. 

<div class="alert alert-info">
**Info**:Use
</div>
<br>
    
```{r code3_11-setup}
dat_female <- .hidden_dat_female
dat_male <- .hidden_dat_male
```
  
```{r code3_11, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_11-setup" }
  
```
  
```{r code3_11-solution}
# Either
fit_f <- lm(sysBP ~ age,data = dat_female)
summary(fit_f)
# Or
fit_m <- lm(sysBP ~ age,data = dat_male)
summary(fit_m)
```
  
```
Results cannot be evaluated automatically - refer to the solution to check your answer.
```
  
</div>
<br>


<div class="box">
#### Exercise 3.12
    
****
    
> Create a scatter plot for the relationship between age and systolic blood pressure in your subsample. 
    
<div class="alert alert-info">
**Info**: Use `age` for the x-axis, and `sysBP` for the y-axis in the `plot()` function. And then use the `fit_f`/`fit_m` object in the `abline()` function.
</div>
<br>
    
```{r code3_12-setup}
dat_female <- .hidden_dat_female
dat_male <- .hidden_dat_male
```
  
```{r code3_12, exercise=TRUE, exercise.eval=FALSE,exercise.setup = "code3_12-setup" }
  
```
  
```{r code3_12-solution}
# Either
plot(dat_male$age, dat_male$sysBP, xlab = "Age", ylab = "Systolic blood pressue")  
abline(fit_m, col="blue",lwd = 3)  
# Or 
plot(dat_female$age, dat_female$sysBP, xlab = "Age", ylab = "Systolic blood pressue")    
abline(fit_f, col="blue",lwd = 3)  
```
  
```
Results cannot be evaluated automatically - refer to the solution to check your answer.
```
  
</div>
<br>



  
<div class="box">
#### Exercise 3.13* - Final exercise
    
****
    
> What is the predicted blood pressue of a 40 year old individual in your subsample? 
    
<div class="alert alert-info">
**Info**: Beta regression coefficients can be retrieved from `fit_m$coefficients`/ `fit_f$coefficients`. 
</div>
<br>
    
```{r code3_13-setup}
dat_female <- .hidden_dat_female
dat_male <- .hidden_dat_male
fit_f <- lm( sysBP ~ age,data = dat_female)
fit_m <- lm( sysBP ~ age,data = dat_male)
```
  
```{r code3_13, exercise=TRUE, exercise.eval=TRUE,exercise.setup = "code3_13-setup" }

```
  
```{r code3_13-solution}
# Either
fit_m$coefficients[1] + 40 * fit_m$coefficients[2] 
# Or 
fit_f$coefficients[1] + 40 * fit_f$coefficients[2] 
```

```
Discuss your result with the group.
```
 

</div>
<br>


****


<!-- create button to link to next page -->
<a href="exercises.Rmd#section-link_Resources" class="btn btn-default skip">Continue to further learning</a>











<!------------------- FURTHER LEARNING ----------------------------------->







## Further learning{#link_Resources}

We hope to see you again on further courses with us at ScHARR. However, alternative resources are available. If you aim to become an expert R user and have a considerable amount of time to learn (for example if you are a PhD student), then we would recommend the following pathway:

### Cheatsheets
**[Cheatsheets](https://rstudio.com/resources/cheatsheets/)** provide quick references for how to use certain R packages or functions. They are not long texts or comprehensive manuals. Instead, they provide visual aids that use layout and visual mnemonics to help people zoom to the functions they need. For a start, you may find the [base R cheatsheet](https://rstudio.com/wp-content/uploads/2016/10/r-cheat-sheet-3.pdat) quite useful.

### Introductory Material
**[R for Data Science](https://r4ds.had.co.nz/)** is a good place to recap the materials taught in this course. The hard copy of Hadley Wickham and Garrett Grolemund's book of the same name (and content) is available on **[amazon](https://r4ds.had.co.nz/)**. A thorough understanding of the content in this book combined with a working knowledge of basic statistics will be enough to undertake independent research on simple topics in most health related fields. 

### Intermediate Material

**[Efficient R Programming](https://csgillespie.github.io/efficientR/)** by Colin Gillespie & Robin Lovelace goes into more detail on optimal workflow, good coding practices, parallelisation and collaboration. Understanding this material is important for running simulations, collaborating with other researchers on projects and working with large datasets. 

### Creating packages

**[R Packages](http://r-pkgs.had.co.nz/)** is a more advanced book, giving an overview of the skills required to build R packages from scratch. This is important in order to move from being a researcher who stands on the shoulders of others, to a researcher on whose shoulders others stand (moving from being a user of packages to a producer of packages).

### Advanced R
**[Advanced R](http://adv-r.had.co.nz/)** describes how R really works. It may be useful for those working on particularly computationally burdensome problems or who are interested in the computing behind R. For most researchers in health this is WAY beyond the level of understanding necessary for the day-job.  

### Swirl
As well as these core texts we also recommend learning R with [swirl](https://swirlstats.com/). Swirl has a range of short courses (approx 30mins) which are undertaken within R. You can download swirl by typing **install.packages("swirl")** into R. Once installed load swirl from the library with **library(swirl)** and then follow instructions within R.

### Data Science for Public Health Economic Evaluation. 
We are in the process of writing a book on data science for public health economic evaluation. It is available online for free here: **[r4phe](https://bookdown.org/rasmith3/r4phe/)**. Once finishes it will cover all the material in this course, then continue on through intermediate material, visualizations, simulations, user interfaces and collaboration. It is designed to aid the training of the next PhD cohort of Public Health Economics & Decision Science at the University of Sheffield.
